Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

quiz 17, STATS102 Fall2019 Session 2, Duke Kunshan University, Dec 2nd
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

Feller discusses the statistics of flying bomb hits in an area in the south of London during the Second World War. The area in question was divided into 24 × 24 = 576 small areas. The total number of hits was 537. Assuming the hits were purely random, use the Poisson approximation to find the probability that a particular square would have exactly k hits. Write a function to compute the probability that one arbitraty area would have k hits. 

<!-- grade_id: myfunction, points: 0, grade: false, solution: true -->
```python
from math import factorial
import numpy as np
def myfunction(k):
    # YOUR CODE HERE
    P=((537/576)**k/factorial(k))*np.exp(-(537/576))
    return P
```

<!-- grade_id: cell-6ae795b2bfe9d088, points: 0.5, grade: true, solution: false -->
```python
assert abs(myfunction(0)-0.39365)<0.01
assert abs(myfunction(1)-0.366997)<0.01
assert abs(myfunction(2)-0.1710)<0.01
assert abs(myfunction(3)-0.05316)<0.01
```