Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

Problem 1 of Assignment 6, STATS102 Fall2019 Session 2, Duke Kunshan University, Due at Nov 7th 12：00pm beijing time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook and nbgrader:

![](jupyter.png)

# Problem1

Given an array of integers where 1 ≤ a[i] ≤ n (n = size of array), some elements may appear twice. Find all the elements of [1, n] inclusive that do not appear in this array. (Return `[]`, when n < 1.)

<!-- grade_id: 111, points: 0, grade: false, solution: true -->
```python
def findDisappearedNumbers(nums):
    # YOUR CODE HERE
    list=[]
    for i in range(1,len(nums)+1):
        if i not in nums:
            list.append(i)
    return list
    if len(nums)==0:
        return []
```

<!-- grade_id: 112, points: 0.5, grade: true, solution: false -->
```python
assert findDisappearedNumbers([4,3,2,7,8,2,3,1,5]) == [6, 9]
assert findDisappearedNumbers([10,9,4,7,8,8,8,1,2,1,1,1,1]) == [3, 5, 6, 11, 12, 13]
assert findDisappearedNumbers([6,7,8,1,2,3,3,3,3,10]) == [4, 5, 9]
assert findDisappearedNumbers([]) == []
```

# Problem 2

Write a program to encode a string(only containing 26 letters) using symbols of international Morse code.

Reference: https://en.wikipedia.org/wiki/Morse_code

For your convienice, you can use the morse code of 26 letters below:

[".-","-...","-.-.","-..",".","..-.","--.","....","..",".---","-.-",".-..","--","-.","---",".--.","--.-",".-.","...","-","..-","...-",".--","-..-","-.--","--.."]

<!-- grade_id: cell-755281e5365d0e99, points: 0, grade: false, solution: true -->
```python
def encodeString(ori_str):
    # YOUR CODE HERE
    list=[".-","-...","-.-.","-..",".","..-.","--.","....","..",".---","-.-",".-..","--","-.","---",".--.","--.-",".-.","...","-","..-","...-",".--","-..-","-.--","--.."]
    letters=['A','B','C','D','E','F','G','H','I','J','K','L','M','N','O','P','Q','R','S','T','U','V','W','X','Y','Z']
    output=''
    for i in ori_str:
        if i in letters:
            output+=list[letters.index(i)]
    return output
```

<!-- grade_id: cell-1c8bbdc679689311, points: 0.5, grade: true, solution: false -->
```python
assert encodeString('CUL') == '-.-...-.-..'
assert encodeString('GA') == '--..-'
assert encodeString('SIG') == '.....--.'
assert encodeString('TU') == '-..-'
```

# Problem 3

Given an input string (s) and a pattern (p), using regular expressions to judge whether s completely matches with p.

Try to think about the result first in your mind.

<!-- grade_id: cell-6aabdfdfdfc07390, points: 0, grade: false, solution: true -->
```python
import re

def isMatch(s, p):
    # YOUR CODE HERE
    return re.search(p,s).span()==(0,len(s))
```

<!-- grade_id: cell-f6a41d77294c95dd, points: 1, grade: true, solution: false -->
```python
assert isMatch("aa","a") == False 
assert isMatch("aa","a*") == True 
assert isMatch("ab",".") == False 
assert isMatch("ab",".*") == True
assert isMatch("aab","a*") == False
assert isMatch("aab",".*") == True
assert isMatch("aab",".*") == True
assert isMatch("aabbcccddd",".*") == True
assert isMatch("aabbcccddd","a*b*.*") == True
assert isMatch("aabbcccddd","a*b*.") == False
assert isMatch("aabbcccddd","a*b*c.d.") == False
assert isMatch("aabbcccddd","a*b*c.*d") == True
```