Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

quiz 13, STATS102 Fall2019 Session 2, Duke Kunshan University, Nov. 20th
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

Load <b>artists.csv</b>, calculate and return how many artists were born in the 19th century [1801-1900].

Data reference: https://www.kaggle.com/ikarus777/best-artworks-of-all-time

<!-- grade_id: myfunction, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def myfunction():
    # YOUR CODE HERE
    data=pd.read_csv('artists.csv')
    data['years']=data['years'].str.replace('–','-')
    data['born']=data['years'].str.extract('(\d+)').astype(int)
    output=np.sum((data['born']>=1801)&(data['born']<=1900))
    
    return output
```

<!-- grade_id: myfunction_test1, points: 0.5, grade: true, solution: false -->
```python
assert myfunction() == 27
```