Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

quiz 10, STATS102 Fall2019 Session 2, Duke Kunshan University,Nov 14th
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

calculate and return the percentage of host country teams ('Country') that won the finals ('Winner') when there are greater or equal to 16 participating teams


\\( \frac{\text{how many times home teams that won the finals when there are greater or equal to 16 participating teams}}{\text{how many world cups when there are greater or equal to 16 participating teams}} \\)

<!-- grade_id: myfunction, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
def myfunction():
    # YOUR CODE HERE
    data= pd.read_csv('WorldCups.csv')
    data1=np.sum(data['QualifiedTeams']>=16)
    data2=np.sum((data['QualifiedTeams']>=16)&(data['Country']==data['Winner']))
    p=(data2/data1)
    return p
```

<!-- grade_id: myfunction_test, points: 0.5, grade: true, solution: false -->
```python
assert abs(myfunction()-0.235)<0.01

```