Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

Problem 1 of Assignment 2, STATS102 Fall2019 Session 2, Duke Kunshan University, Due at Oct 31th 12pm shanghai time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook and nbgrader:

![](jupyter.png)

---
## Part A (1 points)

Write a function that returns a list of numbers, that form a Collatz conjecture sequence. Make sure the input is an int and it handles the case where $n<=0$ by raising a `ValueError`.
https://en.wikipedia.org/wiki/Collatz_conjecture

<!-- grade_id: Collatz, points: 0, grade: false, solution: true -->
```python
def Collatz(n):
    """For input n (must be a positive integer), compute a list of numbers in reverse order, that form a Collatz conjecture sequence.
  
    """
    # YOUR CODE HERE
    if n<=0 or type(n)!=int:
        raise ValueError
    
    seq=[]
    seq.append(n)
    
    while n!=1:
        if n%2==0:
            n=int(n/2)
            seq.append(n)
        
        else:
            n=int(n*3+1)
            seq.append(n)
        
    return seq
    
```

Your function should print `[12, 6, 3, 10, 5, 16, 8, 4, 2, 1]` for $n=12$. Check that it does:

```python
Collatz(12)
```

Output:
```
[12, 6, 3, 10, 5, 16, 8, 4, 2, 1]
```

<!-- grade_id: correct_Collatz, points: 0.5, grade: true, solution: false -->
```python
"""Check that cubics returns the correct output for several inputs"""
assert Collatz(1) == [1]
assert Collatz(2) == [2, 1]
assert Collatz(10) == [10, 5, 16, 8, 4, 2, 1]
assert Collatz(11) == [11, 34, 17, 52, 26, 13, 40, 20, 10, 5, 16, 8, 4, 2, 1]
```

<!-- grade_id: Collatz_invalid_input, points: 0.5, grade: true, solution: false -->
```python
"""Check that cubics raises an error for invalid inputs"""
try:
    Collatz(0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")

try:
    Collatz(-4)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
    
try:
    Collatz(10.0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```

## Part B (1 point)

Write a function to construct a Fibonacci suquence, where the input is a positive number representing the sequence length and output is the list with that length. 

Make sure it handles the case where n is not a positive number by raising a `ValueError`.

Reference: https://en.wikipedia.org/wiki/Fibonacci_number

<!-- grade_id: cell-b7bacc27f96f199c, points: 0, grade: false, solution: true -->
```python
def Fibonacci(n):
    """For input n (must be a positive integer), compute a list of numbers, that form a Fibonacci sequence.
    """
    # YOUR CODE HERE
    if n<=0 or type(n)!=int:
        raise ValueError
        
    list=[]
    if n>0:
        n=int(n)
        if n==1:
            list=[1]
        elif n==2:
            list=[1,1]
        elif n>=3:
            list=[1,1]
            for i in range(3,n+1):
                list.append(list[i-2]+list[i-3])
    return list
    
```

Your function should print `[1, 1, 2, 3, 5, 8, 13, 21, 34, 55]` for $n=10$. Check that it does:

```python
Fibonacci(10)
```

Output:
```
[1, 1, 2, 3, 5, 8, 13, 21, 34, 55]
```

<!-- grade_id: cell-6224cc3d64c1bfd3, points: 0.5, grade: true, solution: false -->
```python
"""Check that Fibonacci returns the correct output for several inputs"""
assert Fibonacci(1) == [1]
assert Fibonacci(3) == [1,1,2]
assert Fibonacci(7) == [1, 1, 2, 3, 5, 8, 13]
assert Fibonacci(15) == [1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610]
```

<!-- grade_id: cell-65bc9effcb936260, points: 0.5, grade: true, solution: false -->
```python
"""Check that Fibonacci raises an error for invalid inputs"""
try:
    Fibonacci(0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise ValueError")

try:
    Fibonacci(-4)
except ValueError:
    pass
else:
    raise AssertionError("did not raise ValueError")
    
try:
    Fibonacci(10.0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise ValueError")
```