Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

Problem 1 of quiz 9, STATS102 Fall2012 Session 2, Duke Kunshan University,Nov 13th
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

Loading `chipotle.tsv`, calculate and return the total revenue of Chipotle Restaurant.

<!-- grade_id: myfunction, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def myfunction():
    chipo = pd.read_csv('chipotle.tsv', sep = '\t')
    # YOUR CODE HERE
    q=np.array(chipo['quantity'])
    p=np.array(chipo['item_price'])
    totalrev=np.sum(q*p)
    return totalrev
```

<!-- grade_id: myfunction_test, points: 0.5, grade: true, solution: false -->
```python
assert (myfunction()-39237.02 < 0.001)
```