Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

Problem 1 of quiz 11, STATS102 Fall2019 Session 2, Duke Kunshan University,Nov 18th
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

Please create and return one DataFrame of name information for Puerto Rico shown as follows:

![](Quiz11.png)

Hint: pd.MultiIndex

<!-- grade_id: myfunction, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def myfunction():
    # YOUR CODE HERE
    
    data={(1998,'Paola'):724,(1998,'Genesis'):500,(2018,'Valentina'):241,(2018,'Victoria'):205}
    df=pd.DataFrame(pd.Series(data))
    df.columns=['Number of occurrences']
    df.index.names=['Year','Name']
    
    return df
```

<!-- grade_id: myfunction_test1, points: 0.5, grade: true, solution: false -->
```python
import numpy as np
import pandas as pd

name_cnt = myfunction()
assert len(name_cnt.index.levels)==2
assert len(name_cnt.index.names)==2
assert len(name_cnt.columns)==1
assert (name_cnt['Number of occurrences'] == [724,500,241,205]).all()
```