Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

Problem 1 of quiz 7, STATS102 Fall2018 Session 1, Duke Kunshan University,Sep11th
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

For an input numpy array x, return a string in the following format:

'the input numpy array has xxx rows, xxx columns!'

xxx should be replaced by the real integer number.

For example:


the input 

array([[0, 1, 2],
       [3, 4, 5],
       [6, 7, 8]])
       
the output

'the input numpy array has 3 rows, 3 columns!'

<!-- grade_id: strictly_ascending_list, points: 0, grade: false, solution: true -->
```python
import numpy as np
def tell_the_shape(x):
    # YOUR CODE HERE
    a=np.shape(x)
    num1=a[0]
    num2=a[1]
    return('the input numpy array has {} rows, {} columns!'.format(num1,num2))
```

<!-- grade_id: strictly_ascending_list1, points: 0.5, grade: true, solution: false -->
```python
import numpy as np
data=np.arange(12).reshape(3,4)
assert tell_the_shape(data)=='the input numpy array has 3 rows, 4 columns!'
data=np.arange(15).reshape(3,5)
assert tell_the_shape(data)=='the input numpy array has 3 rows, 5 columns!'
data=np.empty((4,4))
assert tell_the_shape(data)=='the input numpy array has 4 rows, 4 columns!'
```