Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

Problem 1 of Assignment 4, STATS102 Fall2019 Session 2, Duke Kunshan University, Due at Nov 5th 12：00pm beijing time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook and nbgrader:

![](jupyter.png)

Problem A (1 point)

Write a program to calculate the value of the following piecewise function:

$$f(x)= \begin{cases} -9|x|+16 & \quad x \le 0 \\
            x^2+x+2  & \quad 0 \lt x \lt 21 \\
            x^3-9x & \quad x\ge 21
      \end{cases}$$
  
If the input is not an integer or a float number, then raise ValueError.

<!-- grade_id: 111, points: 0, grade: false, solution: true -->
```python
def myPiecewiseFunction(x):
    # YOUR CODE HERE
    
    if type(x)==int or type(x)==float:
        if x<=0:
            return -9*abs(x)+16
        elif x in range(1,21):
            return x**2+x+2
        else:
            return x**3-9*x
    else:
        raise ValueError
    
```

<!-- grade_id: 112, points: 0.5, grade: true, solution: false -->
```python
assert myPiecewiseFunction(0) == 16
assert myPiecewiseFunction(10) == 112
assert myPiecewiseFunction(21) == 9072
assert myPiecewiseFunction(-10) == -74
assert myPiecewiseFunction(-5) == -29
```

<!-- grade_id: cell-b4b829a6d56dcc50, points: 0.5, grade: true, solution: false -->
```python
try:
    myPiecewiseFunction("5")
except ValueError:
    pass
else:
    raise AssertionError("Error input")
```

Problem B (1 point)

Given a string containing just the characters `'('`, `')'`, `'{'`, `'}'`, `'['` and `']'`, determine whether the input string is complete parenthesis combination. (Empty string is also valid.)

Some correct cases:
`{[()]}`, `[]{}()`, `[()]`, `{}` ...

<!-- grade_id: cell-83d04a5d99c61960, points: 0, grade: false, solution: true -->
```python
def isValid(s):

    # YOUR CODE HERE
    dic = {'{': '}',  '[': ']', '(': ')','?':'?'}
    stack = ["?"]
    for i in s:
        if i in dic: 
            stack.append(i)
        elif dic[stack.pop()] != i: 
            return False 
    return len(stack) == 1
        
```

<!-- grade_id: cell-6947179739c5ffcd, points: 1, grade: true, solution: false -->
```python
"""Check that isValid function returns the correct output"""
assert (isValid("[[[]]{}}") == False)
assert (isValid("[](){}{[]}()[[]]") == True)
assert (isValid("[{}]") == True)
assert (isValid("([)]") == False)
assert (isValid("((((())]") == False)
```