Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

Problem 1 of Assignment 5, STATS102 Fall2019 Session 2, Duke Kunshan University, Due at Nov 6th 12：00pm beijing time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook and nbgrader:

![](jupyter.png)

## Problem1

Sort the input list of tuples [(pay_per_day, working_days),(pay_per_day, working_days)...(pay_per_day, working_days)] in an descending order based on the total pay (per_per_day*working_days), and remove the items when working days are less than 7. 

Try to use sorted function, lambda function and filter function.

<!-- grade_id: 111, points: 0, grade: false, solution: true -->
```python
def sort_total_pay(mylist):
    # YOUR CODE HERE
    new=list(filter(lambda x:x[1]>=7,mylist))
    return sorted(new,key=lambda x:x[0]*x[1],reverse=True)
```

<!-- grade_id: 112, points: 1, grade: true, solution: false -->
```python
assert sort_total_pay([(10.0,10),(15.0,3),(3.1,2),(20.4,2),(4.8,7)])==[(10.0,10), (4.8,7)]
assert sort_total_pay([(15.0,17),(18.0,8),(8.0,3),(10.0,6),(11.0,8)])==[(15.0,17), (18.0,8), (11.0,8)]
```

# Problem 2

Given a string, find all possible positons of a substring in it.   

<!-- grade_id: cell-97ff5d8d38665f39, points: 0, grade: false, solution: true -->
```python
def findSubstr(string, substr):
    
    # YOUR CODE HERE
    list=[]
    for i in range(len(string)):
        if string[i:i+len(substr)]==substr:
            list.append(i)
    return list
```

<!-- grade_id: cell-6a3765707b8aa908, points: 1, grade: true, solution: false -->
```python
assert findSubstr('Hello', 'l')==[2, 3]
assert findSubstr('I love python', 'o')==[3, 11]
assert findSubstr('Banana', 'na')==[2, 4]
assert findSubstr('Good job', 'o')==[1, 2, 6]
```