Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

exercise, STATS102 Fall2019 Session 2, Duke Kunshan University, Nov19th
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

load WorldCupMatches.csv and WorldCupPlayers.csv to calculate all the matches that have goal in the first minute, and output the following dataframe.

![](ex5.png)

<!-- grade_id: myfunction, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
def myfunction():
    # YOUR CODE HERE
    dm=pd.read_csv('WorldCupMatches.csv')
    dp=pd.read_csv('WorldCupPlayers.csv')  
    dm.drop_duplicates(inplace=True)
    dm.dropna(inplace=True)
    dp.drop_duplicates(inplace=True)    #为什么不能合在一起写，一定要分开写？
    
    data=pd.merge(dm,dp,on='MatchID')
    
    data['first goal time']=data['Event'].str.extract('G(\d+)\'.*',expand=False).fillna(0).astype(int)
    
    data1=data.sort_values(by='first goal time',ascending=True)
    data1=data1[data1['first goal time']==1].set_index('MatchID')
    return data1
myfunction()
```

Output:
```
Year              Datetime                  Stage  \
MatchID                                                            
8763.0       1998.0  24 Jun 1998 - 21:00                 Group D   
2347.0       1978.0  02 Jun 1978 - 13:45                 Group 1   
43950063.0   2002.0  29 Jun 2002 - 20:00             Third place   
300186512.0  2014.0  16 Jun 2014 - 19:00                 Group G   
878.0        1982.0  16 Jun 1982 - 17:15                 Group 4   
1146.0       1938.0  05 Jun 1938 - 17:00             First round   
1470.0       1962.0  03 Jun 1962 - 15:00                 Group 4   
1176.0       1938.0  16 Jun 1938 - 18:00             Semi-finals   
1098.0       1930.0  14 Jul 1930 - 14:50                 Group 3   
1544.0       1962.0  07 Jun 1962 - 15:00                 Group 3   
1702.0       1966.0  23 Jul 1966 - 15:00          Quarter-finals   
1105.0       1934.0  07 Jun 1934 - 18:00   Match for third place   

                                 Stadium            City  Home
```

<!-- grade_id: myfunction_test1, points: 0, grade: true, solution: false -->
```python
import numpy as np
import pandas as pd
res=myfunction()
assert res.shape==(12,7)
assert res['Home Team Goals'].sum()==37
assert res['Away Team Goals'].sum()==20
```