Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

Problem 1 of Assignment 15, STATS102 Fall2019 Session 2, Duke Kunshan University, Due at Nov 26th 12：00pm beijing time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook and nbgrader:

![](jupyter.png)

# Problem 1

Let $X_{1}, X_{2}$ be the rolling outcome of two independent fair dies. Let $Y$ denote the minimum of the $X_{1}$ and $X_{2}$. Find the probability of $P(Y=2)$ using both the theory and the simulation.

<!-- grade_id: p1, points: 0, grade: false, solution: true -->
```python
def myfunction1():
    '''theoretical answer, please illustrate how you calculate the answer'''
    # YOUR CODE HERE
    p1=1/6           #p1 is the probability for the die showing 2
    p2=5/6           #p2 is the probability for the die not showing 1
    p=p1*p2*2-p1*p1  #p1*p2 means the probabilty of p1 and p2 happens at the same time,
                     #and we minus p1*p1 because we calculated the probability of both dies showing 2 twice before
    
    return p
```

Output:
```
0.25
```

<!-- grade_id: p2, points: 0, grade: false, solution: true -->
```python
import numpy as np
def myfunction2():
    '''simulation answer, please write codes to simulate the random experiments and calculate the answer'''
    # YOUR CODE HERE
    count=0
    N=1000000
    for i in range(N):
        x1=np.random.randint(1,7)
        x2=np.random.randint(1,7)
        if min(x1,x2)==2:
            count+=1
    return count/N
```

Output:
```
0.250554
```

<!-- grade_id: cell-f6884921eda4f912, points: 1, grade: true, solution: false -->
```python
assert abs(myfunction1()-0.25)<0.01
```

<!-- grade_id: cell-f48e465ce63886b9, points: 1, grade: true, solution: false -->
```python
assert abs(myfunction2()-0.25)<0.01
```