Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

Assignment 10, STATS102 Fall2019 Session 2, Duke Kunshan University, Due at Nov 14th 12：00pm beijing time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook and nbgrader:

![](jupyter.png)

## Problem1
Loading `chipotle.tsv`, calculate the average number of items for each order.

<!-- grade_id: mysin, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def myfunction():
    # YOUR CODE HERE
    chipo=pd.read_csv('chipotle.tsv',sep='\t')
    q=np.sum(chipo['quantity'])
    n=max(chipo['order_id'])
    return q/n    
```

<!-- grade_id: cell-044b32eea7f84461, points: 1, grade: true, solution: false -->
```python
assert (myfunction()-2.711 < 0.01)
```

## Problem2
Loading `chipotle.tsv`, calculate the total income of Chipotle Restaurant.

<!-- grade_id: cell-f20e12447413f5ef, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def myfunction():
    # YOUR CODE HERE
    chipo=pd.read_csv('chipotle.tsv',sep='\t')
    chi=chipo['item_price'].str.replace('$','').astype(float)
    
    totali=np.sum(chipo['quantity']*chi)
    return totali

```

<!-- grade_id: cell-59c863bb5486fd82, points: 1, grade: true, solution: false -->
```python
assert (myfunction()-39237.02 < 0.001)
```