Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

Problem 1 of quiz 16, STATS102 Fall2018 Session 1, Duke Kunshan University,Oct8th
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

A doctor gives a patient a test for a particular cancer. Before the
results of the test, the only evidence the doctor has to go on is that 1 woman
in 100 has this cancer. Experience has shown that, in 99 percent of the cases in
which cancer is present, the test is positive; and in 95 percent of the cases in which
it is not present, it is negative. If the test turns out to be positive, what probability
should the doctor assign to the event that cancer is present?

just show me how you calculate the final result.

<!-- grade_id: myfunction, points: 0, grade: false, solution: true -->
```python
def myfunction():
    # YOUR CODE HERE
    return 
```

<!-- grade_id: cell-df412ed364ca0d56, points: 0, grade: true, solution: false -->
```python
assert abs(myfunction()-1/6)<0.0001
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-17-3beb41fbfd03> in <module>
----> 1 assert abs(myfunction()-1/6)<0.0001

TypeError: unsupported operand type(s) for -: 'NoneType' and 'float'
```