Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

Problem 1 of Assignment 7, STATS102 Fall2019 Session 2, Duke Kunshan University, Due at Nov 8th 12：00pm beijing time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook and nbgrader:

![](jupyter.png)

## Problem1

swap the first two rows of the input matrix and return the resulted matrix

<!-- grade_id: 111, points: 0, grade: false, solution: true -->
```python
import numpy as np
def swap_row12(A):
    # YOUR CODE HERE
    B=np.insert(A,2,A[0],axis=0)
    B=np.delete(B,0,axis=0)
    return B   
```

<!-- grade_id: 112, points: 1, grade: true, solution: false -->
```python
import numpy as np
A = np.arange(25).reshape(5,5)
B=swap_row12(A)
#assert(sum(sum(abs(np.array([[ 5,  6,  7,  8,  9],[ 0,  1,  2,  3,  4],[10, 11, 12, 13, 14],[15, 16, 17, 18, 19],[20, 21, 22, 23, 24]])-B)))==0)
assert((np.array([[ 5,  6,  7,  8,  9],[ 0,  1,  2,  3,  4],[10, 11, 12, 13, 14],[15, 16, 17, 18, 19],[20, 21, 22, 23, 24]])==B).all())
```

# Problem 2

Create a $9*9$ matrix:
$$
 \begin{bmatrix}
   1 & 0 & 3 & 0 & 5 & 0 & 7 & 0 & 9\\
   0 & 2 & 0 & 4 & 0 & 6 & 0 & 8 & 0\\
   11 & 0 & 13 & 0 & 15 & 0 & 17 & 0 & 19\\
   0 & 10 & 0 & 12 & 0 & 14 & 0 & 16 & 0\\
   21 & 0 & 23 & 0 & 25 & 0 & 27 & 0 & 29\\
   0 & 18 & 0 & 20 & 0 & 22 & 0 & 24 & 0\\
   31 & 0 & 33 & 0 & 35 & 0 & 37 & 0 & 39\\
   0 & 26 & 0 & 28 & 0 & 30 & 0 & 32 & 0\\
   41 & 0 & 43 & 0 & 45 & 0 & 47 & 0 & 49\
  \end{bmatrix}
$$

write a function to output this matrix

<!-- grade_id: cell-41d7c4010d18b8dd, points: 0, grade: false, solution: true -->
```python
import numpy as np

def createMatrix():
    # YOUR CODE HERE
    A1=np.arange(1,50,2).reshape(-1,1)
    A1=np.insert(A1,1,0,axis=1).reshape(5,10)
    A1=np.delete(A1,9,axis=1)
    
    A2=np.arange(2,33,2).reshape(-1,1)
    A2=np.insert(A2,0,0,axis=1).reshape(4,8)
    A2=np.insert(A2,8,0,axis=1)
    
    A=np.insert(A1,1,A2[0],axis=0)
    A=np.insert(A,3,A2[1],axis=0)
    A=np.insert(A,5,A2[2],axis=0)
    A=np.insert(A,7,A2[3],axis=0)
    
    return A
    
```

<!-- grade_id: cell-e496b617d982bdd1, points: 1, grade: true, solution: false -->
```python
import numpy as np

assert (createMatrix() == np.array([[ 1,  0,  3,  0,  5,  0,  7,  0,  9],
                                       [ 0,  2,  0,  4,  0,  6,  0,  8,  0],
                                       [11,  0, 13,  0, 15,  0, 17,  0, 19],
                                       [ 0, 10,  0, 12,  0, 14,  0, 16,  0],
                                       [21,  0, 23,  0, 25,  0, 27,  0, 29],
                                       [ 0, 18,  0, 20,  0, 22,  0, 24,  0],
                                       [31,  0, 33,  0, 35,  0, 37,  0, 39],
                                       [ 0, 26,  0, 28,  0, 30,  0, 32,  0],
                                       [41,  0, 43,  0, 45,  0, 47,  0, 49]])).all()
```