Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

Assignment 11, STATS102 Fall2019 Session 2, Duke Kunshan University, Due at Nov 15th 12：00pm beijing time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook and nbgrader:

![](jupyter.png)

## Problem1
Load the file (WorldCups.csv), and then return the following DataFrame sorted by the numbers in the 'Score' column. For each world cup, the winner gets 4 points, the runners-up gets 2 points, the third place gets 1 point.  ('Germany FR' needs to be converted to 'Germany' and count together)

![](assignment9.png)

<!-- grade_id: mysin, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
def myfunction():
    # YOUR CODE HERE
    data=pd.read_csv('WorldCups.csv')
    w=data['Winner'].value_counts()
    r=data['Runners-Up'].value_counts()
    t=data['Third'].value_counts()
    df=pd.DataFrame({'Winner':w,'Runners-Up':r,'Third':t,'Score':0}).fillna(0)
    df['Score']=(4*df['Winner']+2*df['Runners-Up']+df['Third'])
    df.loc['Germany']=df.loc['Germany']+df.loc['Germany FR']
    df=df.drop('Germany FR',axis=0)
    df=df.sort_values(by='Score',ascending=False).astype(int)
    return df
```

<!-- grade_id: cell-044b32eea7f84461, points: 2, grade: true, solution: false -->
```python
import numpy as np
import pandas as pd
res=myfunction()
assert(res.shape==(19, 4))
assert(res['Score'].values==np.array([28, 26, 21, 14,  8,  8,  7,  4,  4,  4,  4,  4,  2,  1,  1,  1,  1,
        1,  1])).all()

```

```python
print('hello world')
```

Output:
```
hello world
```

```python

```