Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

Weekly mini-project 2, STATS102 Fall2019 Session 2, Duke Kunshan University, Due at Nov 11th 12:00pm Shanghai time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

# problem 1 (0.5 point) 
There are 13 lamps on the table. Starting from 1, we turn off the lamp at every 5th, when the count reach 13, start from 1 again. Write a function that return the order of lamps being turned off. In this case, the return list would be [5, 10, 2, 8, 1, 9, 4, 13, 12, 3, 7, 11, 6]. Note that the turn-off lamps would not count again.

<!-- grade_id: cell-06952d5f38104138, points: 0, grade: false, solution: true -->
```python
def turn_off_order():
    # YOUR CODE HERE
    l=[i for i in range(1,14)]
    result=[]
    idx=1
    while len(l)>1:
        idx=idx+4
        while idx>len(l):
            idx-=len(l) #idx = idx - len(l)
        result.append(l[idx-1])
        l.pop(idx-1)
    result.append(l[0])
    return result
    
```

<!-- grade_id: cell-bf13653d315bd398, points: 0.5, grade: true, solution: false -->
```python
assert turn_off_order() == [5, 10, 2, 8, 1, 9, 4, 13, 12, 3, 7, 11, 6]
```

# problem 2 (0.4 point) 
According to problem 1, what if we have different number of lamps, and start at a lamp other than the first lamp. Even more, when we turn off the lamp at every kth. Write a function that aquire 3 numbers, represented the number of lamps, start from which lamp, when to turn off a lamp, and return the turn-off order.

<!-- grade_id: cell-b47caf9d62efd4e1, points: 0, grade: false, solution: true -->
```python
def turn_off_order_upgrade(lamps_num, start, k):
    ''' lamps_num : the number of lamps
       start: the start lamp
       k: turn off the lamp at every k-th lamp
    '''
    # YOUR CODE HERE
    l=[i for i in range(1,lamps_num+1)]
    result=[]
    idx=start
    while len(l)>1:
        idx=idx+(k-1)
        while idx>len(l):
            idx-=len(l) 
        result.append(l[idx-1])
        l.pop(idx-1)
    result.append(l[0])
    return result
```

<!-- grade_id: cell-aac34bc3384ca811, points: 0.4, grade: true, solution: false -->
```python
assert turn_off_order_upgrade(13, 2, 6) == [7, 13, 6, 1, 9, 4, 2, 12, 3, 8, 5, 11, 10]
assert turn_off_order_upgrade(13, 1, 5) == [5, 10, 2, 8, 1, 9, 4, 13, 12, 3, 7, 11, 6]
assert turn_off_order_upgrade(23, 12, 4) == [15, 19, 23, 4, 8, 12, 17, 22, 5, 10, 16, 1, 7, 14, 2, 11, 21, 13, 6, 3, 9, 20, 18]
```

# Problem 3 (0.3 point) 

Find the minimum value of each row for a given 2D matrix.

<!-- grade_id: cell-413652036a3b4661, points: 0, grade: false, solution: true -->
```python
import numpy as np

def findMin(mat):
    # YOUR CODE HERE
    return np.amin(mat,1)
```

<!-- grade_id: cell-4851a061070d469d, points: 0.3, grade: true, solution: false -->
```python
import numpy as np

a = np.arange(1,10).reshape((3,3))
assert (findMin(a) == np.array([1, 4, 7])).all

b = np.linspace(0, 10, 21).reshape((7,3))
assert (findMin(b) == np.array([0., 1.5, 3., 4.5, 6., 7.5, 9.])).all()

c = np.linspace(0, 10, 20).reshape((4,5))
assert ~(findMin(c) == np.array([0., 2.63157895, 5.26315789, 7.89473684])).all()
```

# Problem 4 (0.3 point)

Given two 2D matrices, write a function to judge whether they have the same shape.

<!-- grade_id: cell-6da4d6fe553f2c3a, points: 0, grade: false, solution: true -->
```python
import numpy as np

def judgeShape(A, B):
    # YOUR CODE HERE
    return np.shape(A)==np.shape(B)
```

<!-- grade_id: cell-4b461a8e76c611b7, points: 0.3, grade: true, solution: false -->
```python
import numpy as np

A1 = np.arange(25).reshape(5,5)
B1 = np.arange(20).reshape(4,5)
assert judgeShape(A1, B1) == False

A2 = np.arange(40).reshape(8,5)
B2 = np.arange(8).reshape(8,1)
assert judgeShape(A2, B2) == False

A3 = np.arange(25).reshape(5,5)
B3 = np.arange(0,50,2).reshape(5,5)
assert judgeShape(A3, B3) == True
```

# Problem5 (0.4 point)

Given an array `strList` of strings made only from lowercase letters, return a list of all characters that show up in all strings within the list (no duplicates).

The answer is in sorted order.

<!-- grade_id: cell-82e2f5928a22758d, points: 0, grade: false, solution: true -->
```python
def commonChars(strList):
    # YOUR CODE HERE
    output=[]
    for i in strList[0]:
        if i in strList[1] and i in strList[2] and i not in output:
            output.append(i)
    return sorted(output)
```

<!-- grade_id: cell-447d90122b0aaf79, points: 0.4, grade: true, solution: false -->
```python
assert commonChars(["bella","label","roller"]) == ["e","l"]
assert commonChars(["cool","lock","cook"]) == ['c', 'o']
assert commonChars(["apple","banana","peach"]) == ['a']
assert commonChars(["book","cook","look"]) == ['k', 'o']
```

# Problem6 (0.4 point)

Find the product of the maximum and minimum numbers occuring in `regex_sum_320982.txt` and `regex_sum_42.txt`.

（modified from Python for everyone regular expression assignments）

<!-- grade_id: cell-1cb5eb02bead76e2, points: 0, grade: false, solution: true -->
```python
import re

def find_product(file):
    # YOUR CODE HERE
    alist=[]
    handle=open(file)
    for line in handle:
        num=re.findall('[0-9]+',line)
        alist.extend(num)
    for n in range(len(alist)):
        alist[n]=int(alist[n])
    return max(alist)*min(alist)
```

<!-- grade_id: cell-82183aa329825131, points: 0.4, grade: true, solution: false -->
```python
assert find_product('regex_sum_42.txt') == 29385
assert find_product('regex_sum_320982.txt') == 29682
```

# Problem 7 (0.3 point)

A shift on a string consists of taking the string and moving the leftmost character to the rightmost position. For example, if Str(A) = `abcde`, then it will be `bcdea` after one shift on it.

Given two strings: A and B, please write a function to judge whether A can become B after some number of shifts on A.

<!-- grade_id: cell-d9d13bb04a63a0ec, points: 0, grade: false, solution: true -->
```python
def str_shift(A, B):
    # YOUR CODE HERE
    if len(A)!=len(B):
        return False

    while A!=B:
        a=A[0]
        A=A+a
        A=A[1:]
        if A==B:
            return True
```

<!-- grade_id: cell-de462268a22cf6cd, points: 0.3, grade: true, solution: false -->
```python
# assert str_shift('aa','a') == False
assert str_shift('ab','a') == False
assert str_shift('aeiou','ouaei') == True
assert str_shift('bdffff','ffffbd') == True
```

# Problem 8 (0.4 point)

Write a program to encode the string by spliting it by the middle and reverse the two substrings.

<!-- grade_id: cell-0b95d1b83eb37a8c, points: 0, grade: false, solution: true -->
```python
def encodeString(ori_str):
    # YOUR CODE HERE
    if int(len(ori_str))%2==0:
        
        a1=ori_str[0:int(len(ori_str))//2]
        a1=a1[::-1]
        
        a2=ori_str[int(len(ori_str))//2:]
        a2=a2[::-1]
        
        return a1+a2
     
    if int(len(ori_str))%2!=0:
        
        b1=ori_str[0:(int(len(ori_str))-1)//2]
        b1=b1[::-1]
        
        b2=ori_str[(int(len(ori_str))-1)//2:]
        b2=b2[::-1]
        
        return b1+b2
```

<!-- grade_id: cell-9f5764dd7239884e, points: 0.4, grade: true, solution: false -->
```python
assert encodeString('AaAaBbBb') == 'aAaAbBbB'
assert encodeString('This lines must be printed backwards and splitted in the middle.') \
    == 'wkcab detnirp eb tsum senil sihT.elddim eht ni dettilps dna sdra'
assert encodeString(encodeString('This lines must be printed backwards and splitted in the middle.')) \
    == 'This lines must be printed backwards and splitted in the middle.'
assert encodeString('This is a very secret message!') \
    == ' yrev a si sihT!egassem terces'
assert encodeString(encodeString('This is a very secret message!')) \
    == 'This is a very secret message!'
assert encodeString('Happy programming~') \
    == 'orp yppaH~gnimmarg'
assert encodeString(encodeString('Happy programming~')) \
    == 'Happy programming~'
```