Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

Quiz 12, STATS102 Fall2019 Session 2, Duke Kunshan University, Nov 19th
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

load the census data from http://census.ire.org/data/bulkdata.html (census.csv), calcualte and return a list of top 5 state names ordered in descending way by the average county population in that state. The populations for all counties in each state are listed in column CENSUS2010POP.

<!-- grade_id: myfunction, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
def myfunction():
    # YOUR CODE HERE
    data=pd.DataFrame(pd.read_csv('census.csv'))
    data=data.groupby('STNAME').mean().sort_values(by='CENSUS2010POP',ascending=False) 
    l=[data.index[0],data.index[1],data.index[2],data.index[3],data.index[4]]
    
    #or: list(data.groupby('STNAME')['CENSUS2010POP'].mean().sort_values(ascending=False).index[:5])
    
    return l
```

<!-- grade_id: myfunction_test1, points: 0.5, grade: true, solution: false -->
```python
import numpy as np
import pandas as pd
names=myfunction()
assert names[0]=='California'
assert names[1]=='District of Columbia'
assert names[2]=='Massachusetts'
assert names[3]=='Connecticut'
assert names[4]=='Arizona'
```