Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

Problem 1 of quiz 4, STATS102 Fall2019 Session 2, Duke Kunshan University,Nov 5th
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

Sort the input list of tuples [(pay_per_day, working_days),(pay_per_day, working_days)...(pay_per_day, working_days)] in an descending order based on the total pay (per_per_day*working_days), and return the sorted list. Hint: sorted function with lambda function

<!-- grade_id: strictly_ascending_list, points: 0, grade: false, solution: true -->
```python
def sort_total_pay(mylist):
    # YOUR CODE HERE
    #if mylist==[(10.0,4),(15.0,3),(3.1,2),(20.4,2)]:
        #return [(15.0, 3), (20.4, 2), (10.0, 4), (3.1, 2)]
    
    return sorted (mylist, key=lambda x:x[0]*x[1], reverse=True)
```

<!-- grade_id: strictly_ascending_list1, points: 0.5, grade: true, solution: false -->
```python
assert sort_total_pay([(10.0,4),(15.0,3),(3.1,2),(20.4,2)])==[(15.0, 3), (20.4, 2), (10.0, 4), (3.1, 2)]


```