Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

Lab 3, STATS102 Fall2019 Session 2, Duke Kunshan University, Nov 14th
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

# Problem 1

Loading `'CAvideos.csv`, find the video with the most views and return its title.


<!-- grade_id: strictly_ascending_list, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def most_views():
    # YOUR CODE HERE
    data=pd.read_csv('CAvideos.csv')
    mv=data.sort_values(by='views',ascending=False)
    return mv['title'].iloc[0]
```

<!-- grade_id: strictly_ascending_list1, points: 0, grade: true, solution: false -->
```python
assert most_views() == 'Ed Sheeran - Perfect (Official Music Video)'
```

# Problem 2

Loading `'CAvideos.csv`, find the top 3 video with the most comments and return their titles in a list.

<!-- grade_id: cell-e0009896e4596007, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def top3_comments():
    # YOUR CODE HERE
    data=pd.read_csv('CAvideos.csv')
    mv=data.sort_values(by='comment_count',ascending=False)
    return list(mv['title'].iloc[0:3])
    
```

<!-- grade_id: cell-7debb17a223463d1, points: 0, grade: true, solution: false -->
```python
assert top3_comments() == ['Eminem - Walk On Water (Audio) ft. Beyoncé',
                             'Ed Sheeran - Perfect (Official Music Video)',
                             'Jake Paul - Saturday Night (Song) feat. Nick Crompton & Chad Tepper (Official Music Video)']
```

# Problem 3

Loading `'CAvideos.csv`, find the tag with the highest frequency.

All tags should be lowercase.

<!-- grade_id: cell-be2566487c305f77, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def highest_freq_tag():
    # YOUR CODE HERE
    data=pd.read_csv('CAvideos.csv')
    df=pd.DataFrame(data) 
    tag=df['tags'].str.lower().str.replace('"','').str.split('|')    #a panda series of lists
    a=pd.concat([pd.Series(t) for t in tag])                         #concating multiple series together
    n=a.value_counts()
    return n.index[0]
```

```python
highest_freq_tag()
```

Output:
```
pandas.core.series.Series
```

<!-- grade_id: cell-f8b6f3af9b05ae45, points: 0, grade: true, solution: false -->
```python
assert highest_freq_tag() == 'comedy'
```

# Problem 4

Loading `'CAvideos.csv`, find there are how many videos of which the like ratio is more than $0.9$ and the number of likes is more than $8000$. (Ignore `Nan`) 

$$like\_ratio =  \frac{\text{the number of likes}}{\text{the number of likes + the number of dislikes}} $$

<!-- grade_id: cell-2dd857483e104c7d, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def analyze_likes():
    # YOUR CODE HERE
    data=pd.read_csv('CAvideos.csv')
    data['like_ratio']=data['likes']/(data['likes']+data['dislikes'])
    data=data.dropna()
    a=data.sort_values(by='like_ratio',ascending=False)
    higha=(data['like_ratio']>0.9)&(data['likes']>8000)
    return higha.sum()
```

<!-- grade_id: cell-751309db16b03a89, points: 0, grade: true, solution: false -->
```python
assert analyze_likes() == 96
```