Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Haitong Lin"
COLLABORATORS = ""
```

---

Problem 1 of quiz 7, STATS102 Fall2019 Session 2, Duke Kunshan University,Nov 11th
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

for an input scalar variable x (\\( x \in (-1,1)\\)) and positive integer N, calculate and return \\( \Sigma_{n=0}^{N} (-x)^{n}\\), we will compare the results with \\( \frac{1}{1+x}\\) according to different N. Hint: np.arange np.sum

<!-- grade_id: myfunction, points: 0, grade: false, solution: true -->
```python
import numpy as np
def myfunction(x,N):
    # YOUR CODE HERE
    x=-x
    n=np.arange(0,N+1)
    data=x**n
    A=np.sum(data)
    
    return A
```

<!-- grade_id: myfunction_test, points: 0.5, grade: true, solution: false -->
```python
import numpy as np
assert(abs(myfunction(0.5,10)-1/(1+0.5))>1e-5)
assert(abs(myfunction(0.5,50)-1/(1+0.5))<1e-5)
assert(abs(myfunction(0.5,100)-1/(1+0.5))<1e-10)
assert(abs(myfunction(0.4,5)-1/(1+0.4))>1e-5)
assert(abs(myfunction(0.4,20)-1/(1+0.4))<1e-5)

```